---
title: Test
description: Test

---

## Sample exercise

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
key: 3c871e32e9
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---
## <<<New Exercise>>>

```yaml
type: NormalExercise
key: ff948eb857
lang: python
xp: 100
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---
## <<<New Exercise>>>

```yaml
type: NormalExercise
key: a4642248c1
lang: python
xp: 100
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---
## <<<New Exercise>>>

```yaml
type: NormalExercise
key: b9655a05f0
lang: python
xp: 100
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---
## <<<New Exercise>>>

```yaml
type: NormalExercise
key: 81bc75a1c8
lang: python
xp: 100
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---
## <<<New Exercise>>>

```yaml
type: NormalExercise
key: 9782e3471a
lang: python
xp: 100
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---
## <<<New Exercise>>>

```yaml
type: NormalExercise
key: 10c5f85cc1
lang: python
xp: 100
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---
## <<<New Exercise>>>

```yaml
type: NormalExercise
key: c123b1c33c
lang: python
xp: 100
skills: 2
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
